<problem>

<!-- 
This file is used to make the problem work within the edX system. 
You will not need it if you are using this code for other purposes.
 -->

<script type="text/javascript">
$(document).ready(function(){
	console.log('outer ready');
});

// Logs a javascript object.
function logThatThing(ThatThing){

	// Log it to the console just to verify it's working
	console.log(JSON.stringify(ThatThing));
	
	// Send it to the official edX logamajig!
	Logger.log("harvardx.public_demo.matching_problem", ThatThing);

}

</script>
	
<script type="loncapa/python">

import json

def answercheck(e, ans):
  
  right_answer = []
  right_answer.append( [['A','2'],['B','4'],['C','3'],['D','1'],['D','2']] )
  right_answer.append( [['A','2'],['B','4'],['C','3'],['D','1']] )
  partial_credit = True
  feedback = True

  parsed = json.loads(ans)
  answer = json.loads(parsed['answer'])
  answer = answer['pairings']
  
  if partial_credit:
  
    currentpoints = []
    wrong_answers = []
    maxpoints = []
    scores = []
    answer_index = 0
  
    for right_answer_n in right_answer:

      maxpoints.append(len(right_answer_n))
      currentpoints.append(0)
      wrong_answers.append(0)
      
      # Fix this part.
      # It's the most important part.
      for item in answer:
        does_match = False
        for target in right_answer_n:
          if item == target:
            does_match = True
            break
        if does_match:
          currentpoints[answer_index] += 1
        else:
          wrong_answers[answer_index] += 1
          
    
      scores.append((float(currentpoints[answer_index] - wrong_answers[answer_index])) / float(maxpoints[answer_index]))
      answer_index += 1
  
    final_grade = round(max(scores), 2)
    final_index = scores.index(final_grade)
    message = str(currentpoints[final_index]) + ' correct out of ' + str(maxpoints[final_index]) + ', ' + str(wrong_answers[final_index]) + ' wrong.'
    
    if feedback:
      if final_grade > 0:
        return {
          'input_list': [
            { 'ok': True, 'msg': message, 'grade_decimal': final_grade},
          ]
        }
      else:
        return {
          'input_list': [
            { 'ok': False, 'msg': message, 'grade_decimal': final_grade},
          ]
        }
    else:
      if final_grade > 0:
        return {
          'input_list': [
            { 'ok': True, 'msg': '', 'grade_decimal': final_grade},
          ]
        }
      else:
        return {
          'input_list': [
            { 'ok': False, 'msg': '', 'grade_decimal': final_grade},
          ]
        }
  
  else:
    answer_sort = sorted(answer)
  
    is_right = False
  
    for right_answer_n in right_answer:
      right_answer_sort = sorted(right_answer_n)
  
      if answer_sort == right_answer_sort:
        is_right = True
        break
      else:
        is_right = False
        
    return is_right

       
</script>

<p>This is a matching exercise. Drag the elements on the left into the ones on the right to create a match. You can match items with more than one other item, but you should choose the best match if you can. Click on the items for a better view if you need it.</p>

<customresponse cfn="answercheck">
  <jsinput gradefn="MatchingA.getGrade"
    get_statefn="MatchingA.getState"
    set_statefn="MatchingA.setState"
    width="780"
    height="550"
    html_file="/static/MatchingA.html"
    sop="false"/>
</customresponse>
</problem>