<problem>

<script type="loncapa/python">
<![CDATA[

import json

###################### Set This #########################

grading = 'linear'     # Options: 'linear', 'strict', or 'generous'

############## Leave Everything Else Alone ##############


def answercheck(e, ans):

    # Get the student's answer.
    parsed = json.loads(ans)
    answer = json.loads(parsed['answer'])
    video_length = float(answer['video_length'])
    watch_times = answer['watch_times']
    start_time = float(answer['start_time'])

    durations = []
    total_watch_time = 0
    end_time = 0
    grade = 0

    # Remove duplicates and sort the list
    watch_times = [float(i) for i in watch_times]
    watch_times = list(set(watch_times))
    watch_times.sort()

    # Count up the times
    for j in watch_times:

        ind = watch_times.index(j)
        this_time = j

        next_time = watch_times[ind + 1]

        # If the next one is the last item, mark it as an end and we're done.
        if ind == len(watch_times) - 2:
            end_time = next_time
            durations.append(end_time - start_time)
            break

        # If the next time is too far ahead, call this an end and push duration
        elif next_time - this_time > 2:
            end_time = this_time
            durations.append(end_time - start_time)
            start_time = next_time

        # Otherwise, just keep counting up.
        else:
            pass

    # Add up all the durations to get the total
    total_watch_time = sum(durations)

    grade = total_watch_time / video_length

    if grading == 'strict':
        grade = grade * grade
    elif grading == 'generous':
        grade = math.sqrt(grade)

    # Round up to the nearest tenth.
    grade = math.ceil(grade*10.0) / 10.0

    msg = "You watched about " + str(grade * 100) + " percent of the video."

    if grade > 0.95:
        isOK = True
    elif grade > 0.20:
        isOK = "Partial"
    else:
        isOK = False

    return {
        'input_list': [
            {'ok': isOK, 'msg': '', 'grade_decimal': grade},
        ]
    }

]]>
</script>

<script type="text/javascript">
<![CDATA[

// Logs a javascript object.
function logThatThing(ThatThing){

	// Log it to the console just to verify it's working
	console.log(JSON.stringify(ThatThing));
	
	// Send it to the official edX logamajig!
	Logger.log("harvardx.grape_ape.video_watch_grader", ThatThing);

}

var HXVideoWatch = {
    watch_times: [], 
    video_length: -1,
    start_time: -1
};


$(document).ready(function(){
    // Declaring semi-global variables for later use.
    var video = $('.video');
    var currentTime;
    
    console.log('Video watcher watcher started');
    
    video.each(function(vidnumber){
        
        var thisVid = $(this);
    
        // Check to see whether the video is ready before continuing.
        var waitForVid = setInterval(function(){
            
            try {
                var state = thisVid.data('video-player-state'); // Sometimes this fails and that's ok.

                if(typeof state.videoPlayer !== 'undefined'){
                    
                    if(typeof state.videoPlayer.player.getPlayerState() !== 'undefined'){
                        console.log('video data loaded');
                        HXVideoWatch.start_time = state.videoPlayer.player.getCurrentTime();
                        mainLoop(state, vidnumber);
                        clearInterval(waitForVid);
                    }
                }

            }
            catch(err){
                console.log('waiting for video ' + (vidnumber+1) + ' to be ready');
            }

            
        }, 200);
    
    });
    
    // Every half-second, push the current time to the watch_times array.
    function mainLoop(state, vidnumber){
        
        var timeChecker = setInterval(function(){
            
            // If the video is actually playing...
            if(state.videoPlayer.player.getPlayerState() === 1){
                try{
                    state.videoPlayer.update();        // Forced update of time. Required for Safari.
                }
                catch(err){
                    // If this fails, shut down this loop.
                    // It's probably because we moved to a new tab.
                    clearInterval(timeChecker);
                }
                HXVideoWatch.watch_times.push(state.videoPlayer.player.getCurrentTime());
                
                HXVideoWatch.video_length = state.videoPlayer.player.getDuration();
            }
        }, 500);
    
    }
});

]]>
</script>


<!--    Problem Description Section     -->
<!-- You should ask your question here. -->
<p>Press "Check" when you are done watching the video above.</p>
<!-- End Problem Description Section -->


<!-- You can use this same set of files for every video check problem in your course,
     so you probably don't need to change what's below. -->

<customresponse cfn="answercheck">
  <jsinput gradefn="videowatch.getGrade"
    get_statefn="videowatch.getState"
    set_statefn="videowatch.setState"
    width="800"
    height="0"
    html_file="/static/videowatch.html"
    sop="false"/>
</customresponse>

</problem>