<problem>

  <!--
  --------- Qualtrics Survey Grader Instructions -------
  To use this problem type, you should set:
  1. The survey_length in the python section (optional)
  2. The qualtrics_survey_source in the javascript section (required)
  3. The prompt where it says 'Put your prompt here.' (required)
  4. If applicable, a solution at the end. If not applicable, remove that
     section or your students may see a "show answer" button.
  ----------------- End  Instructions -------------------
  -->

<script type="loncapa/python">
<![CDATA[

from python_lib import HXGraders

def answercheck(e, ans):

    options = {
        'survey_length', 10   # This is the number of pages in the survey.
    }

    # This pulls from the python_lib.zip file in Files & Uploads.
    return HXGraders.qualtricsSurveyGrader(ans, options)

]]>
</script>

<script type="text/javascript">

$(document).ready(function(){
  window.qualtrics_survey_source = 'https://harvard.az1.qualtrics.com/jfe/form/SV_9EvlqyaDbHCPNlz' + '?uid=' + $('#anon_user_id').text();
	console.log('outer ready');
});
function logThatThing(ThatThing){
	console.log(JSON.stringify(ThatThing));
	Logger.log("harvardx.GrapeApe.qualtrics_grader", ThatThing);
}
</script>

<!-- This is to pass some hidden information into the survey. Do not change. -->
<div class="sr" aria-hidden="true">
<span id="anon_user_id">%%USER_ID%%</span>
</div>

<!-- INSTRUCTOR: -->
<p>Put your prompt here.</p>


<!-- Leave this part alone. -->
<p>When you are done with the survey, press "Submit".</p>

<customresponse cfn="answercheck">
  <jsinput title="Graded Survey"
    gradefn="qualtrics_grader.getGrade"
    get_statefn="qualtrics_grader.getState"
    set_statefn="qualtrics_grader.setState"
    width="900"
    height="800"
    html_file="/static/qualtrics_grader.html"
    sop="false"/>
</customresponse>

<!-- Remove this if showing a solution is not appropriate or your purposes. -->
<solution>
<div class="detailed-solution">
<p>Explanation</p>
<p>Questions should have explanations, or at least commentary.</p>
</div>
</solution>


</problem>
